# Homework 3: More String, Lists, and Python Objects (18 pts)

name: Caroline Moy

email: ckmoy@umich.edu

This homework assignment took me 10 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)



## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

I did not discuss this homework with anyone. 

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

```python
import collections
import string
import itertools
```

```python
pip install flake8 pycodestyle_magic
```

Output:
```
Requirement already satisfied: flake8 in /opt/anaconda3/lib/python3.12/site-packages (7.0.0)
Requirement already satisfied: pycodestyle_magic in /opt/anaconda3/lib/python3.12/site-packages (0.5)
Requirement already satisfied: mccabe<0.8.0,>=0.7.0 in /opt/anaconda3/lib/python3.12/site-packages (from flake8) (0.7.0)
Requirement already satisfied: pycodestyle<2.12.0,>=2.11.0 in /opt/anaconda3/lib/python3.12/site-packages (from flake8) (2.11.1)
Requirement already satisfied: pyflakes<3.3.0,>=3.2.0 in /opt/anaconda3/lib/python3.12/site-packages (from flake8) (3.2.0)
Note: you may need to restart the kernel to use updated packages.
```

---
## Question 1: Infinite sequences (8 pts)
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep.

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

<!-- grade_id: primes, points: 0, grade: false, solution: true -->
```python
def primes():
    list_num = [True] * 10000000
    list_num[0] = False
    list_num[1] = False
    index = 2

    while index < len(list_num): 
        if list_num[index]: 
            for i in range(index, len(list_num), index):
                list_num[i] = False
            i += 1
            yield index
        index += 1
    
```

<!-- grade_id: correct_prime, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
p = primes()
plist = [2,3,5,7,11,13,17,19,23,29]
for i in range(10):
    assert next(p) == plist[i]
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

Hint: You are highly encouraged to use relevant functions from itertools to make your solution simple

<!-- grade_id: ruler, points: 0, grade: false, solution: true -->
```python
def ruler():
    
    start_num = 1
    freq_num = 1 # how many times the number has repeated
    freq_list_nums = 1 # how many times that number of numbers has repeated (ex. 2 2 has repeated once, 2233 has repeated twice)
    
    while True: 
       if freq_num == freq_list_nums:
           freq_list_nums = 1
           yield from itertools.repeat(start_num, freq_num)
           start_num += 1
           freq_num += 1
       else:
           yield from itertools.repeat(start_num, freq_num)
           freq_list_nums += 1
           start_num += 1

    
```

<!-- grade_id: correct_ruler, points: 3, grade: true, solution: false -->
```python
r = ruler()
ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
for n in ruler_numbers:
    assert next(r) == n
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.

<!-- grade_id: look_say, points: 0, grade: false, solution: true -->
```python
def look_say(): 
    output = '1'
    yield output

    
    while True: 
        last_output = output
        output_string = ''
        i = 0

        while i < len(last_output):
            count = 1
            while i + 1 < len(last_output) and last_output[i] == last_output[i + 1]:
                count += 1
                i += 1
            output_string = output_string + str(count) + str(last_output[i])
            i += 1
        output = output_string 
    
        yield output 
                
```

<!-- grade_id: correct_look_say, points: 3, grade: true, solution: false -->
```python
s = look_say()
look_say_first_ten = ['1',
 '11',
 '21',
 '1211',
 '111221',
 '312211',
 '13112221',
 '1113213211',
 '31131211131221',
 '13211311123113112211']
for i in range(10):
    s_i = next(s)
    assert isinstance(s_i, str)
    assert s_i == look_say_first_ten[i]
```

### Question 2: Simple ciphers (5 pts)

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

<!-- grade_id: enc_caesar, points: 0, grade: false, solution: true -->
```python
def enc_caesar(s: str, k: int) -> str:
    '''
    Encrypt the message s using Ceasar cipher with key k.
    
    >>> enc_caeasar('thequickbrownfoxjumpsoverthelazydog', 1)
    'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'
    '''
 
    s_list = list(s)
    for i in range(len(s_list)):
        if ord(s_list[i]) + k > 122:
            new_ascii = (ord(s_list[i]) + abs(k) -1 - 122) + 97
            s_list[i] = chr(new_ascii)
        elif ord(s_list[i]) + k < 97:
            new_ascii = 122 - (97 - ord(s_list[i]) + abs(k) - 1)
            s_list[i] = chr(new_ascii)
        else:
            s_list[i] = chr(ord(s_list[i]) + k)
    return str(''.join(s_list))
```

<!-- grade_id: correct_enc_caesar, points: 2, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests:
    assert enc_caesar(msg, k) == enc
```

<!-- grade_id: dec_caesar, points: 0, grade: false, solution: true -->
```python
def dec_caesar(s: str, k: int) -> str:
    '''
    Decrypt the message s using Ceasar cypher with key k.
    
    >>> dec_caeasar('uifrvjdlcspxogpykvnqtpwfsuifmbazeph', 1)
    'thequickbrownfoxjumpsoverthelazydog'
    '''
    # all what's above but backwards 

    s_list = list(s)
    for i in range(len(s_list)):
        if ord(s_list[i]) - k > 122:
            new_ascii = (ord(s_list[i]) + abs(k) -1 - 122) + 97
            s_list[i] = chr(new_ascii)
        elif ord(s_list[i]) - k < 97:
            new_ascii = 122 - (97 - ord(s_list[i]) + abs(k) - 1)
            s_list[i] = chr(new_ascii)
        else:
            s_list[i] = chr(ord(s_list[i]) - k)
    return str(''.join(s_list))
    raise NotImplementedError()
```

<!-- grade_id: correct_dec_caesar, points: 3, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests: 
    assert dec_caesar(enc, k) == msg
```

### Question 3 (5 pts) Valid String.
Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1) Open brackets must be closed by the same type of brackets.
2) Open brackets must be closed in the correct order.
3) Every close bracket has a corresponding open bracket of the same type.


<!-- grade_id: is_valid, points: 0, grade: false, solution: true -->
```python
def is_valid(s):
    open = ['(', '{', '[']
    end = [')', '}', ']']

    begin_dict = {'(':')', '{':'}', '[':']'}
    
    if not len(s) % 2 == 0:
        return False  

    while len(s) > 0:
        if s[0] in end: 
            return False
        first_char = s[0]
        end_char = begin_dict[s[0]]
        if s.find(end_char) == -1:
            return False 
        s = s.replace(first_char, '', 1)
        s = s.replace(end_char, '', 1)
    return True


```

<!-- grade_id: correct_is_valid, points: 5, grade: true, solution: false -->
```python
assert is_valid('()') == True
assert is_valid('([{]}') == False
assert is_valid('([{}])') == True
```